```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

# Lab 08: Exploratory Data Analysis of the Spaceship Titanic Data Set

Welcome to your eighth lab assignment of DATA 271!

### Background/Overview

Kaggle has an ongoing "Getting Started" competition inspired by the original Titanic data set, which we have been using in our demos this week.  More information can be found [here](https://www.kaggle.com/competitions/spaceship-titanic/overview).

The backstory (from Kaggle): The year is 2912, and Earth has received a transmission from four lightyears away and there has been an accident.

The Spaceship Titanic was an interstellar passenger liner launched a month ago. With almost 13,000 passengers on board, the vessel set out on its maiden voyage transporting emigrants from our solar system to three newly habitable exoplanets orbiting nearby stars.

While rounding Alpha Centauri en route to its first destination (the torrid 55 Cancri E) the unwary Spaceship Titanic collided with a spacetime anomaly hidden within a dust cloud. Sadly, it met a similar fate as its namesake from 1000 years before. Though the ship stayed intact, almost half of the passengers were transported to an alternate dimension.

To help rescue crews retrieve the lost passengers, this dataset has been released to challenge people to predict which passengers were transported by the anomaly using records recovered from the spaceship’s damaged computer system.  

In this assignment, we will only complete an exploratory data analysis and we will consider only the training set of data. 

### In today's lab, we will
- Use Seaborn to plot
    - heatmaps 
    - histograms
    - bar plots
    - violin plots
    - KDEs
- Create new columns in a data frame based one existing features
- Address missing values

### 1. Load Packages and Data

**Question 1.1:** Import Numpy, Pandas, MatplotLib, and Seaborn to prepare for the EDA. 

```python

```

```python
grader.check("q1_1")
```

**Question 1.2:** Set the runtime configuration parameters so that the figure size is (6,3) and the default dpi is 80. (Default dpi is 72 in webpages, and we wish to see our figures in higher resolution).

```python
# set defined plot sizes and styles
```

```python
grader.check("q1_2")
```

**Question 1.3:** Import the spaceship titanic data.(It is already in your working directory if you are on JupyterHub. If you are on your personal device, download it from our course repo). 

```python
titanic = ...
titanic.head()
```

```python
grader.check("q1_3")
```

**Question 1.4:** Create a copy of the `titanic` dataframe so that we can check the original data even after we make changes to it. 

```python
titanic_copy = ...
titanic_copy.head()
```

```python
grader.check("q1_4")
```

### 2. Understanding the Data

**Question 2.1:** Use Pandas methods to answer the following questions
- a) What features are included in the data? Put your answer in a list `features_2_1_a`.
- b) What is the shape of the `titanic` dataframe? Put your answer in `shape_2_1_b`.
- c) Indicate if features are categorical or numerical. Put the categoricals in a list `categorical_2_1_c`. Put the numerical variables in `numeric_2_1_c`.

```python
features_2_1_a = ...
features_2_1_a
```

```python
shape_2_1_b = ...
shape_2_1_b
```

```python

categorical_2_1_c = ...
numeric_2_1_c = ...
```

```python
grader.check("q2_1")
```

<!-- BEGIN QUESTION -->

**Question 2.2:** Go to the source of the data on Kaggle and read descriptions of the variables.  Summarize this information using Markdown.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 3. Begin Exploration (Statistics)

**Question 3.1:** How old is the oldest passenger on the ship? Assign `oldest` to a float indicating the age of the oldest passenger. 

```python
oldest = ...
oldest
```

```python
grader.check("q3_1")
```

**Question 3.2:** How old is the youngest passenger on the ship? Assign `youngest` to a float indicating the age of the youngest passenger. 

```python
youngest = ...
youngest
```

```python
grader.check("q3_2")
```

**Question 3.3:** What is the average amount that passengers billed at the food court? Assign `food_court` to a float. 

```python
food_court = ...
food_court
```

```python
grader.check("q3_3")
```

**Question 3.4:** How many cabins are on the ship? Assign `num_cabins` to an int.

```python
num_cabins = ...
num_cabins
```

```python
grader.check("q3_4")
```

**Question 3.5:** How many passengers share a room? Assign `shared_room` to an int.

```python
shared_room = ...
shared_room
```

```python
grader.check("q3_5")
```

### 4. Preprocess the Data

**Question 4.1:** Get the total number of null values per variable. (Your answer should be a Pandas Series).

```python
num_nulls = ...
num_nulls
```

```python
grader.check("q4_1")
```

**Question 4.2:** Visualize the location of the null values with a heatmap.

```python
nullvals_heat = ...
plt.show()
```

```python
grader.check("q4_2")
```

### 5. Visualize the Data

**Question 5.1a:** Create count plots (subplots) for `HomePlanet`, `CryoSleep`, `Destination`, and `VIP`, and discuss your observations.

_Type your answer here, replacing this text._

```python
fig, ax = ...

...
...
...
...

...
```

```python
grader.check("q5_1a")
```

**Question 5.1b:** Create count plots (subplots) for `HomePlanet`, `CryoSleep`, `Destination`, and `VIP`, and bring in the variable of `Transported` into the plots (*HINT*: map it to color).  Discuss any observations you can make about individuals who are transported.

_Type your answer here, replacing this text._

```python
fig2, ax2 = ...

...
...
...
...

...
```

```python
grader.check("q5_1b")
```

**Question 5.2:** Use the plots you created in 5.1a and 5.1b and/or Pandas methods to answer the following questions. 

- Where do most passengers come from?  How many come from that planet? Assign `most_common_home_planet` to a tuple where the first element is the name of the most common home planet and the second is the number of passengers that came from that planet.
- Where are most passengers going?  How many are going to that location? Assign `most_common_destination` to a tuple where the first element is the name of the most common destination and the second is the number of passengers going to the destination.
- Do either home planet or destination seem to relate to an individual being transported or not? Assign `transported_predictor` to a list containing `HomePlanet` or `Destination` or both. 

```python
most_common_home_planet = ...
most_common_destination = ...
transported_predictor = ...
```

```python
grader.check("q5_2")
```

**Question 5.3: Correlations** Create a correlation matrix and a heatmap. Discuss any features which are correlated. Specifically, are there any features correlated with a passenger being transported?

_Type your answer here, replacing this text._

```python
# Create correlation matrix
corrmat = ...
corrmat
```

```python
# plot as a heat map to visualize the information in the matrix
heatmap_fig = ...
...
...
```

```python
grader.check("q5_3")
```

**Question 5.4:** Create a histogram for age. Discuss the qualitative features of the plot and what they mean in terms of the passenger population.

_Type your answer here, replacing this text._

```python
age_dist = ...
...
```

```python
grader.check("q5_4")
```

**Question 5.5a: Age** How does age relate to being transported? Make the following plots and then answer the questions in part b.
- Create a box plot for the age of individuals who were transported and individuals who were not.  Compare the boxplots and five number summary between transported and non-transported individuals.  
- Create a violin plot for the same information.
- Create a histogram for age for both transported and non-transported individuals.
- Create a KDE plot for age for both transported and non-transported individuals.  

```python
# BOXPLOT
...
age_box = ...
...
```

```python
# VIOLIN PLOT
...
age_violin = ...
...
```

```python
# HISTOGRAM
...
age_hist = ...
...
...
age_violin = ...
...
```

```python
# KDE PLOT
...
age_kde = ...
...
```

```python
grader.check("q5_5a")
```

**Question 5.5b:** Select the statements below that align with insights you can gain from the plots you created in part a. 

1. The median age of transported individuals is slightly less than for non-transported individuals.  
2. 0 to 18 year olds are more likely to be transported than not.
3. adults roughly between 18 and 40 are more likely to be transported than not.
4. beyond middle age (~40), there is a pretty equal likelihood of being transported.

Assign `age_conclusions` to a list containing 1, 2, 3, or 4.

```python
age_conclusions = ...
```

```python
grader.check("q5_5b")
```

**Question 5.6a: Spending habits** How do spending habits relate to being transported? Complete the following tasks. These will be used to create more plots in parts b and c.
- Create a copy of `titanic` called `titanic_more_spending_features`.
- Create a new column called `Premium` which shows the total amount the passenger spent on `RoomService`, `Spa` and `VRDeck` combined. 
- Create a new column called `Basic` which shows the total amount the passenger spent on `FoodCourt`, `ShoppingMall` combined.
- Create a new column called `All_Services` which is the total amount the passenger spent. 

```python
titanic_more_spending_features['Premium'] = ...
titanic_more_spending_features['Basic'] = ...
titanic_more_spending_features['All_Services'] = ...
```

```python
grader.check("q5_6a")
```

**Question 5.6b:** Using your new dataframe with additional columns. Recalculate the correlation matrix and plot a heatmap. Do you have any new observations?

_Type your answer here, replacing this text._

```python
plt.figure(figsize = (7,5))
corrmat_update = ...

heatmap_update = ...
...
```

```python
grader.check("q5_6b")
```

**Question 5.6c:** Create a KDE plot for luxury spending (Premium) for both transported and non-transported individuals.  Does this give you any additional information?

_Type your answer here, replacing this text._

```python
plt.figure(figsize = (10,5))
kde_luxury = ...
...
```

```python
grader.check("q5_6c")
```

**Question 5.7a: Family Size** We saw in the 1912 Titanic dataset during lecture that the number of siblings or family members had an impact on survival.  We will investigate if this has an impact on the Spaceship Titanic data.  The data set unfortunately does not have a feature for family size.  However, names are given. Some names are listed as NaN. First, remove those passengers from the dataset. Put your new dataframe with NaN names removed into a copy called `titanic_with_family_size`.

```python
titanic_with_family_size = titanic.copy() 
titanic_with_family_size = ...
```

```python
grader.check("q5_7a")
```

**Question 5.7b:** Split the `Name` feature into a new feature called `LastName`.  This should be a new column in your `titanic_with_lastname` dataframe.

```python
titanic_with_family_size.head()
```

```python
grader.check("q5_7b")
```

**Question 5.7c:** Create a new column in your `titanic_with_lastname` dataframe called `FamilySize` containing the number of individuals in the passengers corresponding family (including the passenger). *Hint:* If you want to split the problem up into steps, you can create a series containing all the unique last names and their counts. Then use the series to populate the `FamilySize` column.

```python
family_size_series = ...
titanic_with_family_size.head()
```

```python
grader.check("q5_7c")
```

**Question 5.7d:** What is the last name of the largest family onboard?

```python
largest_family = ...
largest_family
```

```python
grader.check("q5_7d")
```

**Question 5.7e:** Make a countplot for family size with `Transported` mapped to the color of the bars. What observations can you make? Does family size seem to be an important feature in terms of a passenger being transported or not?

_Type your answer here, replacing this text._

```python
family_size_count = sns.countplot(data = titanic_with_family_size,x='FamilySize',hue='Transported')
plt.show()
```

```python
grader.check("q5_7e")
```

## You're done! 

Congratulations on finishing the lab! Gus is so happy for you! Run the cell below and submit to Canvas. 

<img src="gus_lives_in_tunnel.JPG" alt="drawing" width="500"/>

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> import sys 
>>> 
>>> 'numpy' in sys.modules
True
```
```
>>> 'pandas' in sys.modules
True
```
```
>>> 'matplotlib' in sys.modules
True
```
```
>>> 'seaborn' in sys.modules
True
```


Tests for q1_2 (points: None)
```
>>> plt.rcParams['figure.figsize']
[6.0, 3.0]
```
```
>>> plt.rcParams['figure.dpi']
80.0
```


Tests for q1_3 (points: None)
```
>>> titanic.empty
False
```
```
>>> titanic.shape
(8693, 14)
```


Tests for q1_4 (points: None)
```
>>> titanic_copy.empty
False
```
```
>>> titanic_copy.shape
(8693, 14)
```


Tests for q2_1 (points: None)
```
>>> sum([i in ['PassengerId',
...  'HomePlanet',
...  'CryoSleep',
...  'Cabin',
...  'Destination',
...  'Age',
...  'VIP',
...  'RoomService',
...  'FoodCourt',
...  'ShoppingMall',
...  'Spa',
...  'VRDeck',
...  'Name',
...  'Transported'] for i in features_2_1_a])
14
```
```
>>> shape_2_1_b
(8693, 14)
```
```
>>> sum(i in ['PassengerId', 'HomePlanet','CryoSleep', 'Cabin','Destination','VIP','Name','Transported'] for i in categorical_2_1_c)
8
```
```
>>> sum(i in ['Age','RoomService','FoodCourt','ShoppingMall','Spa','VRDeck']  for i in numeric_2_1_c)
6
```


Tests for q3_1 (points: None)
```
>>> oldest
79.0
```


Tests for q3_2 (points: None)
```
>>> youngest
0.0
```


Tests for q3_3 (points: None)
```
>>> np.round(food_court)
458.0
```


Tests for q3_4 (points: None)
```
>>> num_cabins
6561
```


Tests for q3_5 (points: None)
```
>>> shared_room
3266
```


Tests for q4_1 (points: None)
```
>>> num_nulls
PassengerId       0
HomePlanet      201
CryoSleep       217
Cabin           199
Destination     182
Age             179
VIP             203
RoomService     181
FoodCourt       183
ShoppingMall    208
Spa             183
VRDeck          188
Name            200
Transported       0
dtype: int64
```


Tests for q4_2 (points: None)
```
>>> nullvals_heat.get_xticks()
array([ 0.5,  1.5,  2.5,  3.5,  4.5,  5.5,  6.5,  7.5,  8.5,  9.5, 10.5,
       11.5, 12.5, 13.5])
```
```
>>> nullvals_heat.get_yticks()
array([5.0000e-01, 5.4450e+02, 1.0885e+03, 1.6325e+03, 2.1765e+03,
       2.7205e+03, 3.2645e+03, 3.8085e+03, 4.3525e+03, 4.8965e+03,
       5.4405e+03, 5.9845e+03, 6.5285e+03, 7.0725e+03, 7.6165e+03,
       8.1605e+03])
```


Tests for q5_1a (points: None)
```
>>> ax.size
4
```
```
>>> sum([fig.get_axes()[i].get_xlabel() in ['HomePlanet', 'CryoSleep', 'Destination', 'VIP'] for i in range(4)])
4
```
```
>>> sum([fig.get_axes()[i].get_ylabel() =='count' for i in range(4)])
4
```


Tests for q5_1b (points: None)
```
>>> ax2.size
4
```
```
>>> sum([fig.get_axes()[i].get_xlabel() in ['HomePlanet', 'CryoSleep', 'Destination', 'VIP'] for i in range(4)])
4
```
```
>>> sum([fig.get_axes()[i].get_ylabel() =='count' for i in range(4)])
4
```


Tests for q5_2 (points: None)
```
>>> most_common_home_planet
('Earth', 4602)
```
```
>>> most_common_destination
('TRAPPIST-1e', 5915)
```
```
>>> sum([i in ['Destination', 'HomePlanet'] for i in transported_predictor])
2
```


Tests for q5_3 (points: None)
```
>>> corrmat.shape
(7, 7)
```


Tests for q5_4 (points: None)
```
>>> age_dist.get_xlabel()
'Age'
```


Tests for q5_5a (points: None)
```
>>> age_box.get_xlabel() == 'Transported' or age_box.get_ylabel() == 'Transported'
True
```
```
>>> age_box.get_xlabel() == 'Age' or age_box.get_ylabel() == 'Age'
True
```
```
>>> age_violin.get_xlabel() == 'Transported' or age_violin.get_ylabel() == 'Transported'
True
```
```
>>> age_violin.get_xlabel() == 'Age' or age_violin.get_ylabel() == 'Age'
True
```
```
>>> age_hist.get_xlabel() == 'Age'
True
```
```
>>> age_kde.get_xlabel() == 'Age'
True
```


Tests for q5_5b (points: None)
```
>>> set(age_conclusions)
{1, 2, 4}
```


Tests for q5_6a (points: None)
```
>>> 'Premium' in titanic_more_spending_features.columns
True
```
```
>>> 'Basic' in titanic_more_spending_features.columns
True
```
```
>>> 'All_Services' in titanic_more_spending_features.columns
True
```


Tests for q5_6b (points: None)
```
>>> corrmat_update.shape
(10, 10)
```


Tests for q5_6c (points: None)
```
>>> kde_luxury.get_xlabel()
'Premium'
```


Tests for q5_7a (points: None)
```
>>> titanic_with_family_size.Name.isna().sum()
0
```


Tests for q5_7b (points: None)
```
>>> 'LastName' in titanic_with_family_size.columns
True
```
```
>>> titanic_with_family_size.loc[titanic_with_family_size['Name'] == 'Maham Ofracculy','LastName'] 
0    Ofracculy
Name: LastName, dtype: object
```
```
>>> titanic_with_family_size.loc[titanic_with_family_size['Name'] == 'Willy Santantines','LastName'] 
4    Santantines
Name: LastName, dtype: object
```


Tests for q5_7c (points: None)
```
>>> 'FamilySize' in titanic_with_family_size.columns
True
```
```
>>> titanic_with_family_size.loc[titanic_with_family_size['LastName'] == 'Ofracculy','FamilySize'] 
0    1
Name: FamilySize, dtype: int64
```
```
>>> titanic_with_family_size.loc[titanic_with_family_size['LastName'] == 'Santantines','FamilySize']
4       6
607     6
608     6
1647    6
3793    6
6446    6
Name: FamilySize, dtype: int64
```


Tests for q5_7d (points: None)
```
>>> largest_family
'Casonston'
```


Tests for q5_7e (points: None)
```
>>> family_size_count.get_xlabel()
'FamilySize'
```
