```python
# Initialize Otter
import otter
grader = otter.Notebook("lab10.ipynb")
```

# Lab 10: Web Scraping and Manipulating Data

Welcome to Lab 10 of Data Wrangling and Visualization!

## Overview
The web provides a rich source of data, so we often want to work with that information programmatically in order to make sense of it. Sometimes, that data is provided to us by website creators via csv files or through an API (Application Programming Interface). Other times, we need to collect text from the web ourselves.

It is easy to pull HTML from a website but more difficult to find the information we want from HTML.  Parsing the HTML for targeted information and then storing that information in a structured format will be the focus of this activity. Once we have our data in a structured format, we will practice cleaning and reshaping the data.


## In today's lab, we will
- Become familiar with the structure of HTML and CSS by inspecting source code
- Use the Requests and Beautiful Soup libraries to acquire and parse data from websites
- Work through the challenges of web scraping (i.e. turning messy, unstructured data into workable form)
- Use Pandas methods to clean, manipulate, and reshape data for a brief EDA

This activity will make use of the Requests and Beautiful Soup Python modules.
Documentation can be found here: 
- https://docs.python-requests.org/en/latest/
- https://www.crummy.com/software/BeautifulSoup/

### Import packages


```python
import pandas as pd
import requests
from bs4 import BeautifulSoup
import warnings
warnings.filterwarnings("ignore")
import seaborn as sns
sns.set_style('darkgrid')
import matplotlib.pyplot as plt
import numpy as np
```

### 1. Scrape and Beautify Cal Poly Humboldt Data

**Question 1.1:** Go to https://pine.humboldt.edu/anstud/cgi-bin/filter.pl?relevant=ft1yr_retention.out. Read through the webpage and inspect its source code. What information does this page provide? Once you have a feel for what is on the webpage, import its contents with the requests module.

```python
# Import the data
page = ...
```

```python
grader.check("q1_1")
```

**Question 1.2:** Inspect the text of the data you just imported. 

```python
# inspect the text
page_text = ...
page_text
```

```python
grader.check("q1_2")
```

**Question 1.3:** Create a Beautiful Soup object from `page_text`. 

```python
# Create soup object
soup = ...
soup
```

```python
grader.check("q1_3")
```

<!-- BEGIN QUESTION -->

**Question 1.4:** Use the `prettify()` method to turn the Beautiful Soup parse tree into a nicely formatted string with each HTML tag on its own line. Take a look through it and take note of any tags or attributes you are unfamiliear with. Look up any tags that aren't familiar.

```python
# Print pretty data
```

<!-- END QUESTION -->

### 2. Parse and clean the data

**Question 2.1:** Use a BeautifulSoup method to determine how many tables are on the webpage. 

```python
num_tables = ...
num_tables
```

```python
grader.check("q2_1")
```

**Question 2.2:** Access the second to last table and put the contents in `table1_code`. Inspect it, and then give a brief description of what the table contains. 

_Type your answer here, replacing this text._

```python
table1_code = ...
table1_code
```

```python
grader.check("q2_2")
```

**Question 2.3:** Put the contents from the second to last table (`table1_code`) into Pandas Data Frame.

*HINT 1:* You will probably find it helpful to first create a nested list of the data. Each inner list should contain the contents of one row in the table. The outer list should contain all the rows. Use this nested list to create the Pandas DataFrame.

*HINT 2:* Notice that the first two rows in the table have headers (a slightly different structure from the rest of the table). You might find it helpful to first create the data frame to contain the body of the table. Then create the column names with the headers afterward.

```python
# Create a nested list of data
nested_list = ...


# Put the data in a Pandas dataframe and rename columns as needed
table1_df = ...
table1_df.columns = ...
table1_df
```

```python
grader.check("q2_3")
```

**Question 2.4:** Notice that even rows contain numbers and odd rows contain percents in `table1_df1`. Let's separate this into two Pandas data frames. One should contain only numbers, one should contain percents. 

*NOTE:* Make sure that both tables have data in the `Major Program` column.

```python
# Create numbers data frame
table1_nums = ...
table1_nums
```

```python
# Create percents data frame
table1_percents = ...
table1_percents
```

```python
grader.check("q2_4")
```

<!-- BEGIN QUESTION -->

**Question 2.5:** Check the data types of each column in both tables from the previous question. 

```python
# check nums dtypes
```

```python
# check percents dtypes
```

<!-- END QUESTION -->

**Question 2.6:** Cast the datatypes in your tables appropriately. In other words, if a column contains numbers, cast it to an appropriate numeric data type.

*NOTE:* The percents in the `table1_percents` DataFrame should be floats.

```python
# Cast datatypes in nums table
    
table1_nums.dtypes
```

```python
# Cast datatypes in percents table

table1_percents.dtypes
```

```python
grader.check("q2_6")
```

## 3. Visualize

**Question 3.1:** Visualize the relationship between the 1st Year retention rates (percents) in Fall 13 and the 1st Year retention rates (percents) in Fall 22. Make sure your visualization has a title and your axes are labeled. 

```python
scatter_13_vs_22 = ...
plt.show()
```

```python
grader.check("q3_1")
```

<!-- BEGIN QUESTION -->

**Question 3.2:** Visualize the trend in 1st Year retention rates (percents) over time for Biology. On the same axes plot the trend for English (in a different color). Make sure your plot has an informative legend, axes labels, and title. 

*HINT:* You will probably find it helpful to reshape your `table1_percents` dataframe to create this plot. 

```python
# reshape the data as needed

```

```python
# plot the trends
retention_trends = ...
plt.show()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.3:** Create one more data visualization of your choice. This might require reshaping your data. Describe what insights you gain from your visualization.

```python
# Visualize
```

<!-- END QUESTION -->

## You're done! 

Congratulations on finishing the lab! Gus is proud of you! Run the cell below and submit to Canvas. 

<img src="gus_big_stretch.JPG" alt="drawing" width="500"/>

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> page != ...
True
```
```
>>> page.status_code
200
```


Tests for q1_2 (points: None)
```
>>> page_text != ...
True
```
```
>>> type(page_text) is str
True
```


Tests for q1_3 (points: None)
```
>>> soup != ...
True
```
```
>>> type(soup) is BeautifulSoup
True
```


Tests for q2_1 (points: None)
```
>>> num_tables != ...
True
```
```
>>> num_tables
16
```


Tests for q2_2 (points: None)
```
>>> table1_code != ...
True
```
```
>>> 'First-Time full-time Undergrads by current Major' in table1_code.find('tr').text
True
```


Tests for q2_3 (points: None)
```
>>> type(table1_df) is pd.DataFrame
True
```
```
>>> sum([i in ['Major Program','Fall 13','Fall 14','Fall 15','Fall 16','Fall 17','Fall 18','Fall 19','Fall 20','Fall 21','Fall 22'] for i in list(table1_df.columns)])
11
```
```
>>> table1_df.shape[0]
120
```


Tests for q2_4 (points: None)
```
>>> type(table1_nums) is pd.DataFrame and type(table1_percents) is pd.DataFrame
True
```
```
>>> sum([i in ['Major Program','Fall 13','Fall 14','Fall 15','Fall 16','Fall 17','Fall 18','Fall 19','Fall 20','Fall 21','Fall 22'] for i in list(table1_nums.columns)])
11
```
```
>>> sum([i in ['Major Program','Fall 13','Fall 14','Fall 15','Fall 16','Fall 17','Fall 18','Fall 19','Fall 20','Fall 21','Fall 22'] for i in list(table1_percents.columns)])
11
```
```
>>> table1_nums.shape[0]
60
```
```
>>> "Art" in table1_percents['Major Program'].values
True
```


Tests for q2_6 (points: None)
```
>>> np.array(2).dtype
dtype('int64')
```
```
>>> table1_nums.iloc[:,5].dtype is np.array(2).dtype
True
```
```
>>> table1_percents.iloc[:,5].dtype is np.ones(2).dtype
True
```


Tests for q3_1 (points: None)
```
>>> 'Fall 13'  == scatter_13_vs_22.get_xlabel() or 'Fall 13'  == scatter_13_vs_22.get_ylabel()
True
```
```
>>> 'Fall 22'  == scatter_13_vs_22.get_xlabel() or 'Fall 22'  == scatter_13_vs_22.get_ylabel()
True
```
```
>>> len(scatter_13_vs_22.get_title()) != 0
True
```
