```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Homework 5

Welcome to Homework 5 of DATA 271: Data Wrangling and Visualization! In this assignment, we will practice importing data from an API and exploring data with advanced visualization methods. As always, if you get stuck or have questions feel free to discuss problems with your peers and come to office hours. 

For this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `my_list` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you passed previously. **NOTE:** This homework assignment contains hidden tests. Passing all the auto-graded tests in the assignment does *NOT* mean you have answered everything correctly. So be careful and check your work!

### Import libraries

**Question 0:** Import the appropriate modules for this assignment including NumPy, Pandas, Matplotlib, and Seaborn, Regular Expression (re), and Requests. **(5 Points)**

```python

```

### 1. Importing the data

In this assignment, we will explore housing data from Humboldt county in California. To aquire the data, we will use an API for Zillow. To get started, go to https://www.scrapeak.com/zillow-scraper/ click on "Try for free". Create an account and enter all your relevant information. Once you have created your account and signed in, go to the "Account" tab on the left and copy your API Key at the bottom of the page.

*NOTE:* Creating an account is entirely free. You should **NOT** enter any payment information. If you end up on a page that prompts you for payment info, go to the Plans page and make sure you've selected "Starter."

**Question 1.1:** Assign your copied API Key to the string `api_key`. **(5 Points)**

```python
api_key = ...
```

```python
grader.check("q1_1")
```

**Question 1.2:** Go to https://www.zillow.com/ and in the search bar type "Humboldt County, CA". This should take you to a page with search results. There should be a map on the left side of the page, and a list of properties on the right side. Try using the `For Sale` toggle and switch to look at properties `For Rent`. Click back to `For Sale` since this will be the data we import. Once you get the search results copy the URL (it should be pretty long) and assign it to the variable `zillow_url`. **(5 Points)**

```python
zillow_url = ...
```

```python
grader.check("q1_2")
```

**Question 1.3:** I have created a function called `get_listings` for you. Call the function using your API Key and your Zillow URL as arguments. **(5 Points)**

```python
def get_listings(api_key, url):
    api = "https://app.scrapeak.com/v1/scrapers/zillow/listing"
    return requests.get(api, params={"api_key": api_key,"url":url})

# get listings
listing_response = ...
```

```python
grader.check("q1_3")
```

### 2. Inspecting the data

**Question 2.1:** Check what data type `listing_response` is. **(5 Points)**

```python
listing_response_type = ...
listing_response_type
```

```python
grader.check("q2_1")
```

**Question 2.2:** Recall that a payload in API is the actual data pack that is sent with the `GET` method. The payload can be sent or received in various formats, including JSON.
The two primary parts that make up JSON are keys and values like a dictionary. Use the `.json()` method on the listing response object to get the payload and assign it to `payload`. 

Then inspect the keys of `payload`. 

One of the keys should be called `is_success`. Access the `is_success` key of `payload` and make sure it has the value `True`. **(5 Points)**

```python
payload = ...
```

```python
grader.check("q2_2")
```

**Question 2.3:** Take a look at the value associated with the `data` key in `payload`. What data type is it? **(5 Points)**

```python
response_data = ...
```

```python
grader.check("q2_3")
```

**Question 2.4:** Take some time to inspect the keys of `response_data` to get a feel for what is stored in it. Then, when you are ready, put the data from 
```python
response_data["cat1"]["searchResults"]["mapResults"]
``` 
into a Pandas DataFrame called `humboldt_df`.
**(5 Points)**

```python
humboldt_df = ...
humboldt_df.head()
```

```python
grader.check("q2_4")
```

**Question 2.5:** How many rows and columns are in `humboldt_df`? **(5 Points)**

```python
num_rows = ...
num_cols = ...

num_rows,num_cols
```

```python
grader.check("q2_5")
```

<!-- BEGIN QUESTION -->

**Question 2.6:** To avoid making too many identical API requests in the future, save the data in `humbold_df` to a csv with the `to_csv` method. i.e. run

```python
humboldt_df.to_csv('humboldt_zillow_data.csv') 
```

Check your working directory to verify that you have saved the data in a csv. When you come back to this assignment, you may import the data with `pd.read_csv` instead of using `requests.get` every time. 

**(5 Points)**

```python

```

<!-- END QUESTION -->

**Question 2.7:** Create a copy of `humboldt_df` so that we have an original copy of the data after we make changes to it in our analysis. **(5 Points)**

```python
humboldt_df_copy = ...
```

```python
grader.check("q2_7")
```

**Question 2.8:** Begin inspecting the `humboldt_df` dataframe with methods like `.describe()` and `.info()`. Which columns have null values? Visualize the null values with a heatmap. Comment on what you see. **(5 Points)**

_Type your answer here, replacing this text._

```python

```

```python
null_vals_heat = ...
plt.show()
```

```python
grader.check("q2_8")
```

### 3. Preprocess/Clean the Data

**Question 3.1:** We saw in the previous problem that some of columns have many null values. Drop any columns from `humboldt_df` that have more than 250 null values. **(5 Points)**

```python
humboldt_df = ...
humboldt_df.head()
```

```python
grader.check("q3_1")
```

**Question 3.2:** Take a look at the `price` column. What data type is each element in the column? **(5 Points)**

```python
price_type = ...
```

```python
grader.check("q3_2")
```

**Question 3.3:** We want the price column to be type int. Convert it to type int?

*HINT:* You will probably have to start by removing the dollar signs and the commas from each entry. Then remove the word "From" from the entries that have it. Then replace "K" with "000". Once you have done the steps above, you should be able to convert it to type int. **(5 Points)**

```python
# convert price to int
humboldt_df['price']= ...
humboldt_df.head()
```

```python
grader.check("q3_3")
```

**Question 3.4:** Take a look at the `latLong` column. What datatype is each element of the column? Create two new columns in `humboldt_df`. One should be called `lat` containing the latitude coordinate of the listing and one should be called `lon` containing the longitude coordinate. Once you create these new columns, we will drop the original `latLong` column. **(5 Points)**

```python
# make a lat lon columns
humboldt_df['lat'] = ...
humboldt_df['lon'] = ...
```

```python
# drop latLong
humboldt_df = humboldt_df.drop(columns = 'latLong')
humboldt_df.head()
```

```python
grader.check("q3_4")
```

**Question 3.5:** We want to be able to look at data by city. Create a new column called `city` containing the city that the listing is in. 

*HINT:* The `address` column contains the city before ", CA". Regular expression might help you. **(5 Points)**

```python
# make a lat lon columns
humboldt_df['city'] = ...
humboldt_df.head()
```

```python
grader.check("q3_5")
```

**Question 3.6:** Take a look at the `hdpData` column. What datatype is each element of the column? Using this column, create two new columns in `humboldt_df`. One should be called `homeType` containing the type of listing (e.g. Single Family, Lot/Land, etc.) and one should be called `rentZestimate` containing the rent estimate of the property. 

*NOTE:* There may be some NaN values in the `hdpData`. Keep those values as `NaN` in your new columns. **(5 Points)**

```python
humboldt_df['homeType'] = ...
humboldt_df['rentZestimate'] = ...
humboldt_df.head()
```

```python
grader.check("q3_6")
```

### 4. Visualize the data

**Question 4.1:** Create a plot showing the number of listings in each city. **(5 Points)**

```python
city_counts = ...
```

```python
grader.check("q4_1")
```

**Question 4.2:** Create a boxplot showing the 5 number summary of prices for homes in each city. Comment on what you notice. **(5 Points)**

_Type your answer here, replacing this text._

```python
city_prices = ...
```

```python
grader.check("q4_2")
```

**Question 4.3:** Create a plot showing the count of each home type in the dataset. **(5 Points)**

_Type your answer here, replacing this text._

```python
home_type_counts = ...
```

```python
grader.check("q4_3")
```

**Question 4.4:** Create a violin plot showing the distribution of prices of each home type in the dataset. Comment on what you see. **(5 Points)**

_Type your answer here, replacing this text._

```python
home_type_prices = ...
```

```python
grader.check("q4_4")
```

**Question 4.5:** Create a KDE plot for for the price of each home type. 

*NOTE:* Facet the plot by home type so that you have a separate subplot for each home type. **(5 Points)**

```python
home_type_prices = ...
```

```python
grader.check("q4_5")
```

**Question 4.6:** Create a figure with 6 subplots corresponding to the variables `area`, `beds`, `baths`, `lat`, `lon`, and `price`. Each subplot should contain one boxplot showing the distribution of the variable. Comment on anything interesting you see. **(5 Points)**

_Type your answer here, replacing this text._

```python
boxplots_fig, ax = ...

```

```python
grader.check("q4_6")
```

**Question 4.7:** Create a subset of your dataframe called `single_family` containing only the rows where the `homeType` is `SINGLE_FAMILY`. Then make a joint plot showing the `rentZestimate` on the x-axis and the `price` on the `y-axis`. Map the `beds` variable to color. Comment on what you observe. **(5 Points)**

```python
single_family = ...

single_fam_jointplot = ...
```

```python
grader.check("q4_7")
```

<!-- BEGIN QUESTION -->

**Question 4.8:** Based on the preliminary exploration you have done. Create a research question about this data set and then make a plot that could help you uncover information about your question. For example, one possible question is how does home size vary by location? **(5 Points)**

_Type your answer here, replacing this text._

```python

```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q0 (points: [])


Tests for q1_1 (points: [0])
```
>>> api_key != ...
True
```


Tests for q1_2 (points: [0])
```
>>> zillow_url != ...
True
```


Tests for q1_3 (points: [0])
```
>>> listing_response != ...
True
```


Tests for q2_1 (points: [0])
```
>>> listing_response_type != ...
True
```


Tests for q2_2 (points: [0])
```
>>> payload != ...
True
```


Tests for q2_3 (points: [0])
```
>>> response_data != ...
True
```


Tests for q2_4 (points: [0])
```
>>> (humboldt_df == ...).sum().sum()
0
```


Tests for q2_5 (points: [0])
```
>>> num_rows != ... and num_cols != ...
True
```


Tests for q2_7 (points: [0])
```
>>> (humboldt_df_copy == ...).sum().sum()
0
```


Tests for q2_8 (points: [0])
```
>>> null_vals_heat != ... 
True
```


Tests for q3_1 (points: [0])
```
>>> (humboldt_df == ...).sum().sum()
0
```


Tests for q3_2 (points: [0])
```
>>> price_type != ...
True
```


Tests for q3_3 (points: [0])
```
>>> type(humboldt_df.price[0]) is type(np.ones(2, dtype = 'int')[0])
True
```


Tests for q3_4 (points: [0])
```
>>> 'lat' in humboldt_df.columns
True
```


Tests for q3_5 (points: [0])
```
>>> 'city' in humboldt_df.columns
True
```


Tests for q3_6 (points: [0])
```
>>> 'homeType' in humboldt_df.columns
True
```


Tests for q4_1 (points: [0])
```
>>> city_counts != ...
True
```


Tests for q4_2 (points: [0])
```
>>> city_prices != ...
True
```


Tests for q4_3 (points: [0])
```
>>> home_type_counts != ...
True
```


Tests for q4_4 (points: [0])
```
>>> home_type_counts != ...
True
```


Tests for q4_5 (points: [0])
```
>>> home_type_prices != ...
True
```


Tests for q4_6 (points: [0])
```
>>> boxplots_fig != ...
True
```


Tests for q4_7 (points: [0])
```
>>> single_fam_jointplot != ...
True
```
